# Assignment 3: Chatbots

## Overview

This assignment covers Functions & Algorithms, and builds up to the application of Chatbots. 

This assignment is out of 8 points, worth 8% of your grade.

Make sure your completed assignment passes all the asserts. This assignment also has hidden tests - which means that passing all the asserts you see does not guarantee that you have the correct answer! Make sure to double check and re-run your code to make sure it does what you expect, and that the variables you define end up with values you'd expect. 

## Setup

#### Code Imports

You must run the following cell before proceeding with the last part of the assignment. 

```python
# Setup - run this cell before doing the next part of the assignment
#   This imports some extra code, making it available for us to use later
import string
import random
import nltk
```

### Zen Chatbot

In this assignment, we will be working our way up to building a chatbot.

If you get stuck on anything, come back here, and talk to the Zen chatbot, to relax :)

```python
# Run this cell to talk to a Zan Chatbot
#nltk.chat.zen.demo()
```

## Part 1: Functions 

This part covers the basics of functions.

### Q1 - String Concatenation Function (X points)

Create a function that will concatenate two strings, also adding a space between them. 

To do so, write a function called `string_concatenation` which takes two parameters `string1` and `string2`. 

Inside the function, you are going to concatenate `string1` to `string2`, while also adding a space between the two strings. 

The function should then return the result of this string concatenation.

<!-- grade_id: cell-49ddc5e0ebfe33f2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d8f9d72d0d4c6be5, points: 1, grade: true, solution: false -->
```python
assert callable(string_concatenation)
assert isinstance(string_concatenation('hello', 'world'), str)
assert string_concatenation('hello', 'world') == 'hello world'

```

### Q2 - Find max value  (X points)

Write a function called `find_max`, that takes in a single parameter called `random_list`, which should be a list. 

This function will the max value of the input list, and return it. 

In particular this function will assume list of positive numbers.

To find the max, within the function, use a variable called `list_max`, that you initialize to value 0. 

Then use a loop, to loop across `random_list`. 

Inside the list, use a conditional to check if the current value is larger than `list_max`, and if so, re-assign `list_max` to store this new value.

After the loop, return `list_max`, which should now store the maximum value from within the input list. 

<!-- grade_id: cell-23d051fdc803f0bd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-404d30e1e65b6b83, points: 1, grade: true, solution: false -->
```python
assert callable(find_max)
assert isinstance(find_max([1,7,5]), int)
assert find_max([1,7,5]) == 7

```

### Q3 - Square Everything  (X points)

Write a function called `square_all` that takes an input called `collection` that you assume to be a list or tuple of numbers. 

This function should return a new list, which contains the square of each value in `collection`. 

To do so, you will write a loop inside the function, and also use the `append` method. 

<!-- grade_id: cell-5993fedb7f503a6f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8a86a6f3ec803f9c, points: 1, grade: true, solution: false -->
```python
assert callable(square_all)
assert isinstance(square_all([1, 2]), list)
assert square_all((2, 4)) == [4, 16]
 
```

## Part 2: Refactoring Code to Use Functions

'Refactoring' is the process of updating existing computer code, to keep the same outputs, but with a better internal organization and/or properties. 

We have already written a lot of code in previous assignments. However, since we did not yet have functions, a lot of that code is not very easy to re-use. In this section we will refactor some code - rewriting (and extending) some code from previous assignments as functions, that will be easier to re-use. 

### Q4 - Calculator (X points)

Let's start with the calculator functions that we did in a previous assignment. 

Instead of the conditional, we will make 

Define four functions, called `summation`, `subtraction`, `multiplication`, and `division`. 

Each of them will take two inputs `num1` and `num2`, on use them as follows:
- `summation` will add `num1` to `num2` and return the result
- `subtraction` will subtract `num2` from `num1` and return the result
- `multiplication` will multiply `num1` by `num2` and return the result
- `division` will divide `num1` by `num2` and return the result

<!-- grade_id: cell-6da5de076983020f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2db5b9469ddd3f5e, points: 1, grade: true, solution: false -->
```python
assert callable(summation)
assert isinstance(summation(2, 2), int)
assert summation(1, 1) == 2

```

<!-- grade_id: cell-0e21e32dc2acc953, points: 1, grade: true, solution: false -->
```python
assert callable(subtraction)
assert isinstance(subtraction(2, 2), int)
assert subtraction(1, 1) == 0

```

<!-- grade_id: cell-89646145e140818b, points: 1, grade: true, solution: false -->
```python
assert callable(multiplication)
assert isinstance(multiplication(2, 2), int)
assert multiplication(1, 1) == 1

```

<!-- grade_id: cell-065fccebdd8cda04, points: 1, grade: true, solution: false -->
```python
assert callable(division)
assert isinstance(division(2, 2), float)
assert division(1, 1) == 1

```

### Q4 (cont) - Extending our Calculator  (X points)

Let's now add some more functions to our calculator. 

- Write a function called squared() that takes one parameter, called `num`, and returns the square of it. 
- Write a function called square_root() that takes one parameter, called `num`, and returns the square root of it. 

<!-- grade_id: cell-249d571080483d7e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d9c2b8004a926b36, points: 1, grade: true, solution: false -->
```python
assert callable(square)
assert isinstance(square(2), int)
assert square(5) == 25

```

<!-- grade_id: cell-259646bd6dbbc035, points: 1, grade: true, solution: false -->
```python
assert callable(square_root)
assert isinstance(square_root(4), float)
assert square_root(144) == 12

```

### Example

The goal of using functions is to be able to re-use them, as well as to apply and combine functions together. 

The cell below shows a couple examples of using our calculator functions that we have now defined. 

```python
# Add the multiplication of 2 and 3 to the multiplication of 3 and 4
number1 = multiplication(2, 3)
number2 = multiplication(3, 4)
output1 = summation(number1, number2)
print('Output1: \t', output1)

print('')

# Note that we can call functions as inputs to 
# Therefore, the above construction is equivalent to:
output2 = summation(multiplication(2, 3), multiplication(3, 4))
print('Output2: \t', output2)
```

### Q5 - Using Our Calculator  (X points)

Let's now use our new calculator functions to do some calculations!

- Add the square of 3 to the square root of 9
    - Save the result to a variable called `calc_a`
- Subtract the division of 12 by 3 from the the multiplication of 2 and 4
    - Save the result to a variable called `calc_b`
- Multiply the square root of 16 to the sum of 7 and 9
    - Save the result to a variable called `calc_c`
- Divide the square of 7 by the square root of 49
    - Save the result to a variable called `calc_d`

<!-- grade_id: cell-78166757e331fd3f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e29a42fbdc3ccdc5, points: 1, grade: true, solution: false -->
```python
assert calc_a
assert isinstance(calc_a, float)

```

<!-- grade_id: cell-00d84e7f456e3984, points: 1, grade: true, solution: false -->
```python
assert calc_b
assert isinstance(calc_b, float)

```

<!-- grade_id: cell-c5231a0cfa605a94, points: 1, grade: true, solution: false -->
```python
assert calc_c
assert isinstance(calc_c, float)

```

<!-- grade_id: cell-9b225dabad981fb1, points: 1, grade: true, solution: false -->
```python
assert calc_d
assert isinstance(calc_d, float)

```

### Calling Functions from Functions

In the example above, we were able to refactor our code into re-useable functions. 

We were then able to use these functions and chain them together to flexibly and extensibly do computation.

Next we'll dive into another aspect of writing good modular code: dividing complex tasks into independent functions and then using them together.

In practice, this will look like writing small and specific function to do particular things, and then using those functions in bigger functions.

### Q6 - Functions for Even & Positive Numbers (X points)

Recall Q13 from A2, when we wanted to check a series of conditions across items in a list. 

When writing modular code, we should consider that of these conditions is a separate task, and therefore should be a separate task. 

For question, write two functions, one called `is_even`, and the other `is_positive`. 

Each function should take in a single parameter, a number called `value`. 

- `is_even` will check if that number is even, and return `True` if so, and `False` otherwise. 
- `is_positive` will check if that number is positive, and return `True` if so, and `False` otherwise. 
    - Note

<!-- grade_id: cell-9224aadc2edad42d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-503641a289240788, points: 1, grade: true, solution: false -->
```python
assert callable(is_even)
assert isinstance(is_even(3), bool)
assert is_even(3) == False

```

<!-- grade_id: cell-6fa9bc215ef4255e, points: 1, grade: true, solution: false -->
```python
assert callable(is_positive)
assert isinstance(is_positive(3), bool)
assert is_positive(-2) == False

```

### Q7 - Functions Using Functions (X points)

In this section, your are going to recreate a function called `make_new_list` which takes a list called `data_list` as its input. 

This function will use `is_positive` and `is_even` to mimic the task from A2-Q13. 

That is, this function will  check each value in `data_list` to see if they are even or odd and whether they are positive or negative.
- If they are both positive and even, encode this as the number 1
- If they are both negative and odd, encode this as the the number -1
- Otherwise, encode this as the number 0

- Store the result for each value into a list called `output` and return it.

<!-- grade_id: cell-3566bef39096e64e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3db389fe414cc8f7, points: 1, grade: true, solution: false -->
```python
assert callable(make_new_list)

data_list = [-1, 20, -100, -40, 33, 97, -101, 45, -79, 96]
assert isinstance(make_new_list(data_list), list)
assert(len(make_new_list(data_list)) == len(data_list))

```

## Part 3: Algorithms

This part covers working with algorithms.

### Q8 - Algorithms (X points)

For each of the tasks below, indicate whether you think it is something that a human or a computer would be better at doing. 

- a) xxx
    - Answer in a variable called algo_q_a
- b) xxx
    - Answer in a variable called algo_q_b
- c) xxx
    - Answer in a variable called algo_q_c
- d) xxx
    - Answer in a variable called algo_q_d

<!-- grade_id: cell-af2676bac17c5d61, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b2c5af845ab5e0e3, points: 1, grade: true, solution: false -->
```python
assert algo_q_a in ['human', 'computer']

```

<!-- grade_id: cell-863830055e228463, points: 1, grade: true, solution: false -->
```python
assert algo_q_b in ['human', 'computer']

```

<!-- grade_id: cell-be637c928edced6b, points: 1, grade: true, solution: false -->
```python
assert algo_q_c in ['human', 'computer']

```

<!-- grade_id: cell-6484687fce7ccd33, points: 1, grade: true, solution: false -->
```python
assert algo_q_d in ['human', 'computer']

```

## Part 4:  Chatbots

This part is an application of all the tools we have covered so far, to the problem of creating variations of traditional chatbots. 

A chatbot is a piece of software that tries to have a conversation with a human. While today's chatbots are powerful learners, traditional bots are arguably much simpler. 

Traditional chatbots follow a "rule-based" approach that is defined for them (by you, the programmer). These rules can be very simple to very complex. These types of chatbots can handle simple entries but tend to fail with more complex ones.

To work on this application, we are going to need to use conditions, functions, and algorithims. 


#### Chatbot functions 

In order to make our chatbot, we will create a series of functions that our eventual chatbot will use together to try and chat with a user. At the end, we will test that these functions do indeed allow us to have simple conversations with our chatbot. 

### Example Chatbot - Eliza

The following demo is from the Natural Language ToolKit (nltk), a Python package available through Anaconda. 

Uncomment the line below, and run the code to be able to have a chat with Eliza. 

```python
# Run a demo of a chatbot, with Eliza
#nltk.chat.eliza.demo()
```

### Example Chatbot - Iesha

The following is another demo chatbot from nltk, this one called Iesha. 

As you check out the chatbot demos with Eliza and Iesha, see if you can notice any patterns that might help you figure out how they work (what kind of programming structures they use). 

```python
# Run a demo of a chatbot, with Iesha
#nltk.chat.iesha.demo()
```

### Q9 - Is It a Question (X points)

The first thing we might want to know, given an input to a chatbot, is if the input is a question. 

We are going to use a very simple heuristic to check this: we will consider any input that has a question mark ('?') in it to be a question. 

Write a function called `is_question`, that takes as input a string called `input_string`. 

If there is a '?' in `input_string`, this function should return True, otherwise it should return False. 

<!-- grade_id: cell-06800778424be5ba, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-73224d2951e9985c, points: 1, grade: true, solution: false -->
```python
assert callable(is_question)
assert isinstance(is_question('abc'), bool)
assert is_question('what?') == True

```

### Q10 - Removing punctuation (X points)

Once we know if an input is a question, we are going to get rid of all punctuation. 

To do so, write a function called `remove_punctuation`, which takes an input called `input_string`. 

This function should:
- define a new variable `out_string` as an empty string
- loop through each character in `input_string`
    - if this character is not in `string.punctuation`
        - then append the current character to `out_string`
- return `out_string`

<!-- grade_id: cell-e89fcaa0b5c38b5b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7653654ef6ba58e1, points: 1, grade: true, solution: false -->
```python
assert callable(remove_punctuation)
assert isinstance(remove_punctuation('a'), str)
assert remove_punctuation("Hey! It's Tom!") == "Hey Its Tom"

```

### Q11 - Preparing Text (X points)

Now we want a more general function to prepare out text data for processing. 

Write a function called `prepare_text`, that takes as input `input_string`. 

Inside this function, using string methods and our `remove_punctuation` function, this function should:
- make the string all lower case
- remove all punctuation from the string
- split the string into words (look for a string method called `split`)

<!-- grade_id: cell-c1c23cc422a6e514, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-617b5f8e65de0792, points: 1, grade: true, solution: false -->
```python
assert callable(prepare_text)
assert isinstance(prepare_text('One two three.'), list)
assert prepare_text('Hi! Also, howdy.') == ['hi', 'also', 'howdy']


```

### Q12 - Echo a Response (X points)

Define a function, called `respond_echo`, that will return a string a specified number of times, with a given separator. 

Input
* `input_string` - a string
* `number_of_echoes` - an integer
* `spacer` - a string

Output
* `echo_output` - string

This function should return the `input_string`, repeated the number of times specified by `number_of_echoes`, joined by the  `spacer`. 

So, for example, `respond_echo('bark', 2, '-')` should return: 'bark-bark-'

Note that a simple version of this version will end up appending the spacer to the end of the list as well. For our purposes here, implement that simple version.

<!-- grade_id: cell-d77bc7ed64441ebc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3280829d10a28901, points: 1, grade: true, solution: false -->
```python
assert callable(respond_echo)
assert isinstance(respond_echo('ha', 2, ' '), str)
assert respond_echo('meow', 3, '~') == 'meow~meow~meow~'

```

### Q13 - Selector (X points)

Next we will write a function that will help us select an output for the chatbot, based on the input it got. 

Define a function, called `selector`, with the following inputs and outputs. 

Input
* `input_list` - a string
* `check_list` - an integer
* `return_list` - a string

Output
* `output` - can be a string, or None

Inside this function:
- Initialize `output` to `None`
- Loop through each element in input_list
- Use a conditional to check if this element is in `check_list`
    - If it is, assign `output` as the returned value of calling the `random.choice` function on `return_list`
    - Also, break out of the loop
- At the end of the function, return `output`

<!-- grade_id: cell-a85973a26e490194, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-88e6e54e20d32647, points: 1, grade: true, solution: false -->
```python
assert callable(selector)

assert selector(['in', 'words'], ['words'], ['yes']) == 'yes'
assert selector(['in', 'words'], ['out'], ['yes']) == None

```

### Q14 - List to String (X Points)

Since we are often dealing with a list of strings, it would be useful to have a function to turn a list of string back into one single concatenated string. 

Write a function called `list_to_string` that takes two agruments: a list called `input_list` and a string called `seperator`. 

This function will return a string that is each element of `input_list` concatenated together with each sperated by the string `seperator`. 

For example, the following function call:
`list_to_string(['This', 'is', 'fun'], ' ')`

shoud return: 'This is fun.'

<!-- grade_id: cell-cfabdfe3966db8c6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b8d8c7506f53317e, points: 1, grade: true, solution: false -->
```python
assert callable(list_to_string)
assert isinstance(list_to_string(['a', 'b'], '|'), str)
assert list_to_string(['a', 'b'], '|') == 'a|b|'

```

### Q15 - Multi String Concatenator (X points)

Define a function, `multi_concat`, that takes three string inputs, `str_1`, `str_2` and `str_3. 

This function should return the three string concatenated together, with spaces beteween them. 

<!-- grade_id: cell-9c45fdf1f810cb7e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-081f2946d90cf799, points: 1, grade: true, solution: false -->
```python
assert callable(multi_concat)
assert isinstance(multi_concat('a', 'b', 'c'), str)
assert multi_concat('a', 'b', 'c') == 'a b c'

```

### Q16 - End Chat (X points)

The last thing we need is a way to end the chat with our chatbot. 

We will use the same convention that `nltk` uses - namely, that we will end the conversation if the word `quit` appears.

To be able to do so, write a function called `end_chat`, that will take as input a variable called `input_list`, and return `output`, which will be a boolean. 

`end_chat` should return True if the string 'quit' is in the list of strings, `input_list`, and return False otherwise. 

<!-- grade_id: cell-1ab514a0dcaee24e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7937207e3e365a07, points: 1, grade: true, solution: false -->
```python
assert callable(end_chat)
assert isinstance(end_chat(['a', 'b']), bool)
assert end_chat(['quit']) == True

```

```python
# This cell defines a collection of input and output things our chatbot can say and respond to

GREETINGS_IN = ['hello', 'hi', 'hey', 'hola', 'welcome', 'bonjour', 'greetings']
GREETINGS_OUT = ["Hello, it's nice to talk to you!", 'Nice to meet you!',  "Hey - Let's chat!"]

COMP_IN = ['python', 'code', 'turing', 'computer']
COMP_OUT = ["Python is my mother.", \
            "Did you know I'm made of code!?", \
            "Computers are so magical", \
            "Do you think I'll pass the Turing test?"]

JOKES_IN = ['funny', 'hilarious', 'ha', 'haha', 'hahaha', 'lol']
JOKES_OUT = ['ha', 'haha', 'lol'] 

NONO = ['matlab', 'java', 'C++']
NONO_EXT = "I'm sorry, I cannot talk about "

QUESTION = "I'm too shy to answer questions. What do you want to talk about?"
UNKNOWN = ['Good.', 'Okay', 'Huh?', 'Yeah!', 'Thanks!']
```

## Talking to our Chatbot

If everything has gone well with writing all the pieces of our chatbot above, the cell below should let us chat to our bot!

You don't have to write anything in the code block below. 

Do read through it though, to see how our chatbot works, and to see how all the functions you just wrote are used to create the chatbot.

```python
# Run this cell

chat = True
while chat:
    
    # Get a message from the user
    msg = input('INPUT :\t')
    out_msg = None
    
    # Check if the input is a question
    question = is_question(msg)
    
    # Prepare the input message
    msg = prepare_text(msg)
    
    # Check for an end msg 
    if end_chat(msg):
        out_msg = 'Bye!'
        chat = False
    
    # Check for a selection of topics that we have defined to respond to
    if not out_msg:
        out_1 = selector(msg, GREETINGS_IN, GREETINGS_OUT)
        out_2 = selector(msg, COMP_IN, COMP_OUT)
        out_3 = respond_echo(selector(msg, JOKES_IN, JOKES_OUT), 3, '')

        options = list(filter(None, [out_1, out_2, out_3]))
        if options:
            out_msg = random.choice(options)
    
    # If we don't have an output yet, but the input was a question, return msg related to it being a question
    if not out_msg and question:
        out_msg = QUESTION

    # Catch-all to say something if msg not caught & processed so far
    if not out_msg:
        out_msg = random.choice(UNKNOWN)

    print('OUTPUT:', out_msg)
```

After you've explored chatting to the chatbot, you are done!

## The End!

This is the end of the assignment!
Have a look back over your answers, and also make sure to Restart & Run All from the kernel menu to double check that everything is working properly. 

When you are ready to submit your assignment, upload it to TritonED under Assignment-3.